# Working with multiple data sets

There are two data files that we'll be working with for this week's assignment.  They are described below.  Load those data files in with Pandas and then work to answering each of the questions below.  All of these files are found in our usual s3 bucket: `https://hds5210-data.s3.amazonaws.com`

## npidata.csv

This file is basic information about every healthcare provider in the US.  It has one row for each NPI (National Provider Identifier).  It contains information such as the provider's name and address.

* https://hds5210-data.s3.amazonaws.com/npidata.csv


## cmsYYYY.csv

These are files about what kinds of procedures and patients providers in the US are serving under CMS programs, Medicare and Medicaid.  Each contains various statistics about providers over the course of a year.  There are three of these, for the years 2014, 2015, and 2016.  These files, however, don't contain information about the provider such as where the provider is located.

* https://hds5210-data.s3.amazonaws.com/cms2014.csv
* https://hds5210-data.s3.amazonaws.com/cms2015.csv
* https://hds5210-data.s3.amazonaws.com/cms2016.csv

## Our Goals

For this assignment, we're going to want to compute some statistics based on the data in the **cms** files, but aggregate that data based on information in the **npidata** file.  As in last week's assignment, you'll need to store your answers in a variable called `answer` at the end of each step.

```python
import pandas as pd
```

## Part 1

In this first step, we'll need to merge together all of the **cms** files into a single dataframe.  Be careful that these files might not be identical, so you'll have to look a little bit to figure out how to merge them.

As you are merging them, make sure that you retain information about which file (i.e. which year) the data came from.  Call that new columns `year`.

In your `answer` variable, provide a complete data frame that contains all of the rows and columns from the **cms** files, plus an additional column to store the year/file that particular row came from.

The assertion tests will give you a good idea as to if you're merging the files correctly.

<!-- grade_id: 01-solution, points: 0, grade: false, solution: true -->
```python
answer = None

### SOLUTION
CMS_2014 = pd.read_csv('https://hds5210-data.s3.amazonaws.com/cms2014.csv')
CMS_2015 = pd.read_csv('https://hds5210-data.s3.amazonaws.com/cms2015.csv')
CMS_2016 = pd.read_csv('https://hds5210-data.s3.amazonaws.com/cms2016.csv')

CMS_2014['year'] = 2014
CMS_2015['year'] = 2015
CMS_2016['year'] = 2016

CMS_2014.columns = CMS_2014.columns.str.lower()
CMS_2015.columns = CMS_2015.columns.str.lower()
CMS_2016.columns = CMS_2016.columns.str.lower()

CMS_2014.columns = CMS_2014.columns.str.replace(' ', '_')
CMS_2015.columns = CMS_2015.columns.str.replace(' ', '_')
CMS_2016.columns = CMS_2016.columns.str.replace(' ', '_')



answer = pd.concat([CMS_2014, CMS_2015, CMS_2016])

answer.head()


```

Output:
```
nbr         npi                    provider_type  \
0    1  1003000134                        Pathology   
1   19  1003001249               Physical Therapist   
2   21  1003001322  Mass Immunization Roster Biller   
3   37  1003002486                     Chiropractic   
4   73  1003005430            Clinical Psychologist   

  medicare_participation_indicator  number_of_hcpcs  total_services  \
0                                Y               13          8643.0   
1                                Y                4          1221.0   
2                                Y                7           788.0   
3                                Y                2           142.0   
4                                Y                3           130.0   

   total_unique_benes  total_submitted_chrg_amt  total_medicare_allowed_amt  \
0                4276                1357139.00                   321641.88   
1                  69                  88686.28                    34060.09   
2       
```

<!-- grade_id: 01-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (193862, 60))
assert(list(answer['year'].unique()) == [2014,2015,2016])
assert(set(answer.columns.str.lower()) == set(['year', 'nbr', 'npi', 'provider_type',
       'medicare_participation_indicator', 'number_of_hcpcs', 'total_services',
       'total_unique_benes', 'total_submitted_chrg_amt',
       'total_medicare_allowed_amt', 'total_medicare_payment_amt',
       'total_medicare_stnd_amt', 'drug_suppress_indicator',
       'number_of_drug_hcpcs', 'total_drug_services',
       'total_drug_unique_benes', 'total_drug_submitted_chrg_amt',
       'total_drug_medicare_allowed_amt', 'total_drug_medicare_payment_amt',
       'total_drug_medicare_stnd_amt', 'med_suppress_indicator',
       'number_of_med_hcpcs', 'total_med_services', 'total_med_unique_benes',
       'total_med_submitted_chrg_amt', 'total_med_medicare_allowed_amt',
       'total_med_medicare_payment_amt', 'total_med_medicare_stnd_amt',
       'beneficiary_average_age', 'beneficiary_age_less_65_count',
       'beneficiary_age_65_74_count', 'beneficiary_age_75_84_count',
       'beneficiary_age_greater_84_count', 'beneficiary_female_count',
       'beneficiary_male_count', 'beneficiary_race_white_count',
       'beneficiary_race_black_count', 'beneficiary_race_api_count',
       'beneficiary_race_hispanic_count', 'beneficiary_race_natind_count',
       'beneficiary_race_other_count', 'beneficiary_nondual_count',
       'beneficiary_dual_count', 'beneficiary_cc_afib_percent',
       'beneficiary_cc_alzrdsd_percent', 'beneficiary_cc_asthma_percent',
       'beneficiary_cc_cancer_percent', 'beneficiary_cc_chf_percent',
       'beneficiary_cc_ckd_percent', 'beneficiary_cc_copd_percent',
       'beneficiary_cc_depr_percent', 'beneficiary_cc_diab_percent',
       'beneficiary_cc_hyperl_percent', 'beneficiary_cc_hypert_percent',
       'beneficiary_cc_ihd_percent', 'beneficiary_cc_ost_percent',
       'beneficiary_cc_raoa_percent', 'beneficiary_cc_schiot_percent',
       'beneficiary_cc_strk_percent', 'beneficiary_average_risk_score']))
```

```python

```

## Part 2

In this next part, we're going to join the **cms** data with the provider information in the **/data/npidata.csv** file.  In this join, we don't want to lose any records from the **cms** files, even if no matching provider exists in the **npidata** file.  However, we don't care about any providers from the **npidata** file that don't have records in the **cms** files.  Those providers can be ignored.

Join the data files together to create one unified data frame called `answer`.  This dataframe should have all the columns from both **cms** and **npidata** files, joined together using the `npi` column.  

Note that `npi` is unique in the **npidata** file.

<!-- grade_id: 02-solution, points: 0, grade: false, solution: true -->
```python
answer = None

### SOLUTION
import pandas as pd
import numpy as np

# Reading the input data
npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')

# Merging CMS data from different years
answer = pd.concat([CMS_2014, CMS_2015, CMS_2016])

# Renaming columns for consistency
answer = answer.rename(columns={'npi': 'NPI'})

# Specifying the common columns and relevant columns from npidata
common_cols = ['NPI']
npidata_cols = ['Provider Business Mailing Address State Name']

# Performing the first merge with just necessary columns
merged_data = pd.merge(answer, npidata[common_cols + npidata_cols], how='left', on='NPI')

# Updating answer with the merged data
answer = merged_data[answer.columns.tolist() + npidata_cols]

# Merging any other missing columns from npidata
for col in npidata.columns:
    if col not in answer.columns and col != "NPI":
        answer = pd.merge(answer, npidata[['NPI', col]], on='NPI', how='left')



# Final shape output
answer.shape

```

Output:
```
<ipython-input-4-a1b7531500dd>:8: DtypeWarning: Columns (18,33) have mixed types. Specify dtype option on import or set low_memory=False.
  npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')
(193862, 102)
```

<!-- grade_id: 02-tests, points: 5, grade: true, solution: false -->
```python
import numpy
assert(answer.shape == (193862, 103))
assert(list(answer['Provider Business Mailing Address State Name'].unique()) == ['IL','MO',numpy.nan,'WY'])
assert(list(answer.groupby('Provider Business Mailing Address State Name').NPI.count()) == [111520, 53366, 4805])
```

## Part 3

If you did everything right above, you'll notice that grouping by a column with NaN in it will cause some rows to disappear from the aggregation test.  So, let's create a new column called `'State'` that has the same value as whatever is in the `'Provider Business Mailing Address State Name'` column or a value of `'XX'` if there is no state information.

Set `answer` to be your final data frame with the new `'State'` column added.

<!-- grade_id: 03-solution, points: 0, grade: false, solution: true -->
```python
answer = None

### SOLUTION
import pandas as pd
import numpy as np

# Reading the input data
npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')

# Merging CMS data from different years
answer = pd.concat([CMS_2014, CMS_2015, CMS_2016])

# Renaming columns for consistency
answer = answer.rename(columns={'npi': 'NPI'})

# Specifying the common columns and relevant columns from npidata
common_cols = ['NPI']
npidata_cols = ['Provider Business Mailing Address State Name']

# Performing the first merge with just necessary columns
merged_data = pd.merge(answer, npidata[common_cols + npidata_cols], how='left', on='NPI')

# Updating answer with the merged data
answer = merged_data[answer.columns.tolist() + npidata_cols]

# Merging any other missing columns from npidata
for col in npidata.columns:
    if col not in answer.columns and col != "NPI":
        answer = pd.merge(answer, npidata[['NPI', col]], on='NPI', how='left')
# Ensuring 'Provider Other Organization Name' is included.
if 'Provider Other Organization Name' not in answer.columns:
    answer = pd.merge(answer, npidata[['NPI', 'Provider Other Organization Name']], on='NPI', how='left')

# Creating the 'State' column, replacing NaN values in 'Provider Business Mailing Address State Name' with 'XX'
answer['State'] = answer['Provider Business Mailing Address State Name'].fillna('XX')


answer.shape


```

Output:
```
<ipython-input-7-022a61ec24d4>:8: DtypeWarning: Columns (18,33) have mixed types. Specify dtype option on import or set low_memory=False.
  npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')
(193862, 103)
```

<!-- grade_id: 03-tests, points: 5, grade: true, solution: false -->
```python
assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
assert(answer.shape == (193862, 104))
```

Output:
```
<ipython-input-11-ead5e4353622> in <cell line: 1>()
----> 1 assert(list(answer.groupby('State').npi.count()) == [111520, 53366, 4805, 24171])
      2 assert(answer.shape == (193862, 104))

/usr/local/lib/python3.10/dist-packages/pandas/core/groupby/groupby.py in __getattr__(self, attr)
   1361             return self[attr]
   1362 
-> 1363         raise AttributeError(
   1364             f"'{type(self).__name__}' object has no attribute '{attr}'"
   1365         )

AttributeError: 'DataFrameGroupBy' object has no attribute 'npi'
```

## Part 4

Next, let's summarize the data by year and by State.  Create a pivot table that contains one row for each state and one column for each year.  Within the pivot table, put a sum of total services as the values.

Assign `answer` to be that resulting pivot table.  In the tests, I'm going to plot a bar chart of your pivot table.

```python
answer = None
# SOLUTION
import pandas as pd
import numpy as np

# Reading the input data
npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')

# Merging CMS data from different years
answer = pd.concat([CMS_2014, CMS_2015, CMS_2016])

# Renaming columns for consistency
answer = answer.rename(columns={'npi': 'NPI'})

# Specifying the common columns and relevant columns from npidata
common_cols = ['NPI']
npidata_cols = ['Provider Business Mailing Address State Name']

# Performing the first merge with just necessary columns
merged_data = pd.merge(answer, npidata[common_cols + npidata_cols], how='left', on='NPI')

# Updating answer with the merged data
answer = merged_data[answer.columns.tolist() + npidata_cols]

# Merging any other missing columns from npidata
for col in npidata.columns:
    if col not in answer.columns and col != "NPI":
        answer = pd.merge(answer, npidata[['NPI', col]], on='NPI', how='left')
# Ensuring 'Provider Other Organization Name' is included.
if 'Provider Other Organization Name' not in answer.columns:
    answer = pd.merge(answer, npidata[['NPI', 'Provider Other Organization Name']], on='NPI', how='left')

# Creating the 'State' column, replacing NaN values in 'Provider Business Mailing Address State Name' with 'XX'
answer['State'] = answer['Provider Business Mailing Address State Name'].fillna('XX')
# Printing the shape to help debug
print(answer.shape)
answer.head()

answer = pd.pivot_table(answer, values='total_services', index='State', columns='year', aggfunc=np.sum) # now 'answer' is assigned to a dataframe instead of None
answer
```

Output:
```
<ipython-input-12-f8264dc6c533>:7: DtypeWarning: Columns (18,33) have mixed types. Specify dtype option on import or set low_memory=False.
  npidata = pd.read_csv('https://hds5210-data.s3.amazonaws.com/npidata.csv')
(193862, 103)
<ipython-input-12-f8264dc6c533>:39: FutureWarning: The provided callable <function sum at 0x7c2de8700670> is currently using DataFrameGroupBy.sum. In a future version of pandas, the provided callable will be used directly. To keep current behavior pass the string "sum" instead.
  answer = pd.pivot_table(answer, values='total_services', index='State', columns='year', aggfunc=np.sum) # now 'answer' is assigned to a dataframe instead of None
year          2014         2015         2016
State                                       
IL     101393758.4  103048716.3  105432032.4
MO      48577227.1   50808468.0   50854632.1
WY       3487240.6    3800978.9    3604487.9
XX      14938636.4   16533705.2   16705781.4
```

<!-- grade_id: 04-tests, points: 5, grade: true, solution: false -->
```python
assert(answer.shape == (4,3))
assert(answer.sum().sum() == 519185664.6999999)
assert(answer[2016].sum() == 176596933.80000004)
assert(answer.loc['WY'].sum() == 10892707.4)
```

Output:
```
AssertionError                            Traceback (most recent call last)
<ipython-input-13-db10226393b0> in <cell line: 2>()
      1 assert(answer.shape == (4,3))
----> 2 assert(answer.sum().sum() == 519185664.6999999)
      3 assert(answer[2016].sum() == 176596933.80000004)
      4 assert(answer.loc['WY'].sum() == 10892707.4)

AssertionError:
```

```python
%matplotlib inline
answer.plot.bar()
```

```python

```